Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise 1 - Shell basics

Work through as much of the Software Carpentry [lesson on the Unix Shell](http://swcarpentry.github.io/shell-novice/) as you can.  Run through the Setup section just below, then open a terminal through Jupyter to run through the exercises.

After you have completed the first few sections of the tutorial, return to this notebook.  Execute all of the cells, and answer all of the questions.


## Setup - getting required files 

To get started, you'll need to have the required files in your directory.  Use `wget` to get them:

```python
!wget http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
```

```python
!unzip -l shell-novice-data.zip
```

```python
!unzip shell-novice-data.zip
```

*Note*: you only need to do this once per session while using Jupyter on datanotebook.org.  You can open a terminal now and work through the steps, and return to this notebook a little later, and the files will be available either way.  That's because you're working in the same local directory.

**However**!  If you download this file, close your Jupyter session for the night, then come back tomorrow and open up a new Jupyter session on the server again, you'll need to get those sample files again.  Just execute the cells above to do it.

Okay, let's get on with the exercise!

## Navigating Files and Directories

As you work through this section of the tutorial, complete the steps here as well, using the `!` shell escape command.  Execute each cell as you go.

These steps aren't exactly the same as what's in the tutorial, where the file layout is a little different and where they're not using a notebook like we are.  That's okay.  Just consider this practice.

```python
!whoami
```

```python
!pwd
```

```python
!ls -F
```

```python
!ls -F
```

```python
!ls -F data-shell/
```

```python
!ls -aF
```

```python
!ls -af .
```

What is the difference between the two previous cells, and what does the single dot mean?

<!-- grade_id: 156, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
!ls -F ..
```

What do the double dots mean?

<!-- grade_id: 155, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
!ls data-shell/north-pacific-gyre/2012-07-03/
```

## Working with Files and Directories

The following cells come from the next section of the tutorial.

```python
!ls -F
```

```python
!mkdir thesis
```

<!-- grade_id: 121, points: 1, grade: true, solution: false -->
```python
import os
assert "thesis" in os.listdir()
```

```python
!ls -F
```

You can't use the nano editor here in Jupyter, so we'll use the `touch` command to create an empty file instead.

```python
!touch thesis/draft.txt
```

<!-- grade_id: 123, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir("thesis")
```

```python
!ls -F thesis
```

Removing files and directories.

```python
!rm thesis/draft.txt
```

<!-- grade_id: 125, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" not in os.listdir("thesis")
```

```python
!rm thesis
```

```python
!rmdir thesis
```

<!-- grade_id: 127, points: 1, grade: true, solution: false -->
```python
assert "thesis" not in os.listdir()
```

```python
!ls
```

Renaming and copying files.

```python
!touch draft.txt
```

<!-- grade_id: 129, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir()
```

```python
!mv draft.txt quotes.txt
```

<!-- grade_id: 130, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "draft.txt" not in os.listdir()
```

```python
!ls
```

```python
!cp quotes.txt quotations.txt
```

<!-- grade_id: 134, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "quotations.txt" in os.listdir()
```